Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "XAVIER MALTAS TARRIDAS"
COLLABORATORS = ""
```

---

# ![Spark Logo](http://spark-mooc.github.io/web-assets/images/ta_Spark-logo-small.png)  ![Python Logo](http://spark-mooc.github.io/web-assets/images/python-logo-master-v3-TM-flattened_small.png)
# Comptant paraules: Construeix una aplicació que compti paraules de manera eficient

Amb l'ús massiu d'Internet i les xarxes socials, el volum de text no estructurat aquesta creixent dramàticament, i Spark és una gran eina per a analitzar aquest tipus de dades. En aquesta activitat, escriurem codi per a trobar les paraules més comunes en un text generat, el ja conegut [Lorem Ipsum](https://www.lipsum.com/).

El més interessant de la manera de treballar en aquesta pràctica és que podria escalar-se per a analitzar contingut web, per exemple, dins de Wikipedia.


## Durant aquesta activitat cobrirem:

* Part 1:* Creació d'un RDD i un pair RDD (1 Punt)
* Part 2:* Comptar paraules fent servir un pair RDD (1 Punt)
* Part 3:* Trobar paraules concretes i la seva freqüència d'aparició mitjana (1 Punt)
* Part 4:* Aplicar les funcionalitats desenvolupades a un arxiu de text (2 Punts)
* Part 5:* Calcular algunes estadístiques (2 Punts)
* Part 6:* Vídeo resposta (3 Punts)

> Com a referència a tots mètodes que es requereixen per a implementar aquesta pràctica podeu consultar:
> * [API Python de Spark](https://spark.apache.org/docs/2.4.0/rdd-programming-guide.html)

## Part 1: Creació d'un RDD i un pair RDDs

En aquesta secció, explorarem com crear RRDs fent servir 'parallelize' i com aplicar pair RDDs al problema del comptatge de paraules.

### (0a) Què és un RDD (Resilient Distributed Dataset)?

Com hem vist, un RDD és una col·lecció d'elements sobre el qual podem aplicar operacions que es poden ejectuar en paral·lel, amb les següents característiques:

- Són immutables
- Estan distribuïts en els diferents nodes del clúster
- Són tolerants a fallades

### (0b) Configuració de l'entorn python + spark
> Com a referència a tots mètodes que es requereixen per a implementar aquesta pràctica podeu consultar:
> * [API Python de Spark](https://spark.apache.org/docs/2.4.0/rdd-programming-guide.html)

```python
import findspark
findspark.init()

from pyspark import SparkConf, SparkContext

conf = SparkConf()
conf.setMaster("local[1]")

```

Output:
```
<pyspark.conf.SparkConf at 0x7f90189505c0>
```

Introduïu el nom de l'app ActividadRDDs_ seguit del vostre nom d'usuari

Copieu aquestes líneas a la cel·la seguent amb el vostre usuari.

    conf.setAppName("ActividadRDDs_usuario")
    sc = SparkContext.getOrCreate(conf=conf)
    sc.version

<!-- grade_id: cell-9832a75750bbeb02, points: 0, grade: true, solution: true -->
```python
conf.setAppName("ActividadRDDs_xmaltast")
sc = SparkContext.getOrCreate(conf=conf)
sc.version
```

Output:
```
'2.4.0-cdh6.2.0'
```

### (1a) Creació d'un RDD
Comencem generant un RDD a partir d'una llista de Python i el mètode 'sc.parallelize'. Després mostrarem per pantalla el tipus de la variable generada.

```python
wordsList = ['lion', 'cat', 'tiger', 'elephant', 'rat', 'tiger', 'rat', 'cat']
wordsRDD = sc.parallelize(wordsList, 4)
# Print out the type of wordsRDD
type(wordsRDD)
```

Output:
```
pyspark.rdd.RDD
```

### (1b) Crear el plural de les paraules i testejar

Utilitzarem una transformació 'map()' per a incorporar la lletra 's' a cadascun dels strings emmagatzemats en el RDD que acabem de crear. Primer definirem una funció de Python que retorni la paraula que se li ha passat com a paràmetre incorporant una "s" al final d'aquesta.

En aquest exercici, es demana reemplaçar el text `<FILL IN>` amb la solució proposada. Després d'haver definit correctament la funció 'makePlural', executar la següent cel·la que conté una sèrie d'asserts de test. Si la solució és correcta, s'imprimirà '1 test passed'.

Aquesta serà la forma habitual de treballar en les Activitats. Els exercicis contindran una explicació del que s'espera, seguit d'una cel·la de codi amb un o més `<FILL IN>`. Les cel·les que necessitin ser modificades contindran el text '# TODO: Replace with appropriate code' en la primera línia. En alguns casos, com veureu al final de la descripció de l'enunciat tindreu una petita ajuda (HINT) per a completar l'exercici.

Una vegada s'hagin substituït tots els `<FILL IN>` pel codi Python adequat, executar la cel·la, i posteriorment executar la cel·la següent de test per a comprovar que la solució és l'esperada.

<!-- grade_id: cell-b8b4b950bec37953, points: 0, grade: false, solution: true -->
```python
# TODO: Replace <FILL IN> with appropriate code
def makePlural(word):
    return word + 's'

# Apply the makePlural function to each element in the RDD using map()
pluralRDD = wordsRDD.map(makePlural)

# Collect the results to the driver program and display them
pluralWords = pluralRDD.collect()

# Display the resulting plural words
print(pluralWords)
```

Output:
```
['lions', 'cats', 'tigers', 'elephants', 'rats', 'tigers', 'rats', 'cats']
```

<!-- grade_id: cell-babaf7e3cda35be8, points: 0.2, grade: true, solution: false -->
```python
# TEST Pluralize and test (1b)
assert makePlural('lion')=='lions'
assert makePlural('cat')=='cats'
assert makePlural('tiger')=='tigers'
```

### (1c) Aplicar 'makePlural' al nostre RDD
Ara és el moment d'aplicar la nostra funció 'makePlural()' a tots els elements del RDD usant la transformació [map()](http://spark.apache.org/docs/2.4.0/api/python/pyspark.html#pyspark.RDD.map). Es demana completar la següent cel·la i posteriorment executar l'acció [collect()](http://spark.apache.org/docs/2.4.0/api/python/pyspark.html#pyspark.RDD.collect) en la cel·la de test per a obtenir el RDD transformat.

A continuació, s'ofereix una petita ajuda.

**HINT (Copy & Paste and Complete):**
'''python
pluralRDD = wordsRDD.map(< FILL_IN >)
'''

<!-- grade_id: cell-e72735b916b19a0d, points: 0, grade: false, solution: true -->
```python
pluralRDD = wordsRDD.map(makePlural)
```

<!-- grade_id: cell-8f6820c332a87cf7, points: 0.2, grade: true, solution: false -->
```python
# TEST Apply makePlural to the base RDD(1c)
assert pluralRDD.collect() == ['lions', 'cats', 'tigers', 'elephants', 'rats', 'tigers', 'rats', 'cats']
```

### (1d) Aplicar una funció 'lambda' al llarg dels valors en un RDD mitjançant la funció 'map'

Ara realitzarem una transformació equivalent a l'anterior mitjançant una funció 'lambda' en lloc d'una funció amb nom.

Es demana completar la següent cel·la amb la solució. Podeu basar-vos en el següent HINT, i els exemples que es mostren en [API Python de Spark](https://spark.apache.org/docs/2.4.0/rdd-programming-guide.html)

**HINT (Copy & Paste and Complete):** 
```python
pluralLambdaRDD = wordsRDD.map(< FILL IN >)
```

<!-- grade_id: cell-4565359139bf05c4, points: 0, grade: false, solution: true -->
```python
# lambda function to pluralize words
pluralLambdaRDD = wordsRDD.map(lambda word: word + 's')
```

<!-- grade_id: cell-92dcd0dadb60af2e, points: 0.2, grade: true, solution: false -->
```python
# TEST Cell: run collect action to the RDD (pluralLambdaRDD) obtained in the previous cell (1d)
assert pluralLambdaRDD.collect() == ['lions', 'cats', 'tigers', 'elephants', 'rats', 'tigers', 'rats', 'cats']
```

### (1e) Nombre de caràcters de cadascuna de les paraules

Ara es demana obtenir el nombre de caràcters de cada paraula mitjançant un transformació 'map()' i una funció 'lambda'. Farem servir l'acció 'collect' per a guardar el resultat directament en una variable.

**HINT (Copy & Paste and Complete):** 
```python
pluralLengths = pluralRDD.map(< FILL_IN >).< FILL_IN >
```

<!-- grade_id: cell-9b0af843e7c83e09, points: 0, grade: false, solution: true -->
```python
# lambda function to show number of character in each word
pluralLengths = pluralRDD.map(lambda word: len(word)).collect()

# Display character count per each word
print(pluralLengths)
```

Output:
```
[5, 4, 6, 9, 4, 6, 4, 4]
```

<!-- grade_id: cell-362b1d5aa87c3abd, points: 0.2, grade: true, solution: false -->
```python
# TEST Length of each word (exercice 1e)
assert pluralLengths == [5, 4, 6, 9, 4, 6, 4, 4]
```

### (1f) Pair RDDs

Per a completar el nostre programa de comptatge de paraules ara treballarem amb els anomenats pair RDD. Un pair RDD és un RDD en el qual cada element és un tupla de l'estil `(k, v)` on `k` és la clau i `v` és el seu valor corresponent. Primerament, es demana crear un pair RDD que contingui tuplas amb el format `(<word>, 1)` per a cada element del nostre RDD bàsic.

Podem crear nostre pair RDD usant una transformació `map()` i una funció 'lambda()' que transformi adequadament els elements del RDD base.

**HINT (Copy & Paste and Complete):** 
```python
wordPairs = wordsRDD.map(< FILL IN >)
```

<!-- grade_id: cell-de89e345a5dccc45, points: 0, grade: false, solution: true -->
```python
#lambda function to create pair RDDs
wordPairs = wordsRDD.map(lambda word: (word, 1))

#Display pair RDDs
print(wordPairs.collect())
```

Output:
```
[('lion', 1), ('cat', 1), ('tiger', 1), ('elephant', 1), ('rat', 1), ('tiger', 1), ('rat', 1), ('cat', 1)]
```

<!-- grade_id: cell-d79816693c5197a0, points: 0.2, grade: true, solution: false -->
```python
# TEST Pair RDDs (1f)
assert wordPairs.collect() == [('lion', 1), ('cat', 1), ('tiger', 1), ('elephant', 1), ('rat', 1), ('tiger', 1), ('rat', 1), ('cat', 1)]
```

## Part 2: Comptar paraules fent servir un pair RDD


Ara, es demana comptar el nombre de vegades que cada paraula en particular apareix en el RDD. Aquesta operació es pot realitzar d'una infinitat de maneres, però algunes seran molt menys eficients que unes altres.

Una solució molt senzilla seria usar 'collect()' sobre tots els elements retornar-los al driver i una vegada recollits al driver contar-los. Mentre aquesta manera de treballar podria funcionar amb textos relativament curts, nosaltres el que volem és poder treballar amb textos de qualsevol longitud. Addicionalment, executar tot el càlcul en el driver és molt més lent que executar-lo en paral·lel en els workers. Per aquests motius, en aquesta pràctica farem servir operacions paral·lelitzables.

### (2a) Fent servir 'groupByKey()'
Una primera proposta de solució al nostre problema, després veurem que hi ha unes altres molt més eficients, es podria basar en la transformació [groupByKey()](http://spark.apache.org/docs/2.4.0/api/python/pyspark.html#pyspark.RDD.groupByKey). Com el seu nom indica, la transformació 'groupByKey()' agrupa tots els elements d'un RDD que comparteixin la mateixa clau en una única llista dins d'una de les particions.

Aquesta operació planteja dos problemes:
+ Necessita moure tots els valors dins de la partició adequada. Això satura la xarxa.
+ Les llistes generades poden arribar a ser molt grans arribant fins i tot a saturar la memòria d'algun dels workers

Veurem com solucionar aquests problemes posteriorment. De moment es demana usar transformació  'groupByKey()' per a generar un pair RDD del tipus '('word', iterator)'.

**HINT (Copy & Paste and Complete):**
```python
wordsGrouped = wordPairs.< FILL IN >
for key, value in wordsGrouped.collect():
    print('{0}: {1}'.format(key, list(value)))
```

<!-- grade_id: cell-baba5ba9cccf54d0, points: 0, grade: false, solution: true -->
```python
# Use groupByKey to group words by their keys
wordsGrouped = wordPairs.groupByKey()

# Iterate through the groups and print the word and its occurrences
for key, value in wordsGrouped.collect():
    print('{0}: {1}'.format(key, list(value)))
```

Output:
```
lion: [1]
tiger: [1, 1]
cat: [1, 1]
elephant: [1]
rat: [1, 1]
```

<!-- grade_id: cell-bf339be44df0b0dc, points: 0.25, grade: true, solution: false -->
```python
# TEST groupByKey() approach (2a)
assert sorted(wordsGrouped.mapValues(lambda x: list(x)).collect()) == [('cat', [1, 1]), ('elephant', [1]), ('lion', [1]), ('rat', [1, 1]), ('tiger', [1, 1])]
```

### (2b) Utilitza 'groupByKey()' per a obtenir els contadors

Fent servir la transformació 'groupByKey()' crea un RDD que contingui 2 elements, on cadascun d'ells sigui una tupla paraula (clau) iterador de Python (valor). Es pot reaprofitar el RDD 'wordsGrouped'.

Després suma tots els valors del iterador usant una transformació 'map()'. El resultat ha de ser un pair RDD que contingui les parelles (word, count).

**HINT (Copy & Paste and Complete):** 
```python
wordCountsGrouped = wordsGrouped.< FILL IN >
```

<!-- grade_id: cell-c7dea62d6c0c55ca, points: 0, grade: false, solution: true -->
```python
wordCountsGrouped = wordsGrouped.map(lambda x: (x[0], sum(x[1])))

#Print the result
for (word, count) in wordCountsGrouped.collect():
    print('{}: {}'.format(word, count))
```

Output:
```
lion: 1
tiger: 2
cat: 2
elephant: 1
rat: 2
```

<!-- grade_id: cell-74a69c127898b698, points: 0.25, grade: true, solution: false -->
```python
# TEST Use groupByKey() to obtain the counts (2b)
assert sorted(wordCountsGrouped.collect())==[('cat', 2), ('elephant', 1), ('lion', 1), ('rat', 2), ('tiger', 2)]
```

### (2c) Contadors fent servir 'reduceByKey'

Intentant solucionar la problemàtica citada anteriorment, una millor solució per a obtenir els comptatges és començar des d'un pair RDD i després usar la transformació [reduceByKey()](http://spark.apache.org/docs/2.4.0/api/python/pyspark.html#pyspark.RDD.reduceByKey) per a crear un nou pair RDD. La transformació 'reduceByKey()' agrupa totes les parelles que comparteixen la mateixa clau. Posteriorment, aplica la funció que se li passa per paràmetre agrupant els valors de dos en dos. Aquest procés es repeteix iterativament fins que aconseguim un únic valor agregat per a cadascuna de les claus del pair RDD. La transformació 'reduceByKey()' opera aplicant la funció primer dins de cadascuna de les particions de manera independent, i posteriorment només comparteix els valors agregats entre particions diferents, permetent-li escalar de manera eficient donat que no té necessitat de desplaçar per la xarxa una gran quantitat de dades.

**HINT (Copy & Paste and Complete):** 
```python
wordCounts = wordPairs.reduceByKey(< FILL IN >)
```

<!-- grade_id: cell-a2dfc0df90101b5b, points: 0, grade: false, solution: true -->
```python
wordCounts = wordPairs.reduceByKey(lambda x, y: x + y)

#Print the 
for (word, count) in wordCounts.collect():
    print('{}: {}'.format(word, count))
```

Output:
```
lion: 1
tiger: 2
cat: 2
elephant: 1
rat: 2
```

<!-- grade_id: cell-4336283b7de4f325, points: 0.25, grade: true, solution: false -->
```python
# TEST Counting using reduceByKey (2c)
assert sorted(wordCounts.collect())==[('cat', 2), ('elephant', 1), ('lion', 1), ('rat', 2), ('tiger', 2)]
```

### (2d) Ara tot junt

La versió més complexa per a solucionar el problema del comptatge executa primer un 'map()' sobre el pair RDD, després la transformació 'reduceByKey()' i, finalment, l'acció 'collect()' en una única línia de codi.

**HINT (Copy & Paste and Complete):** 
```python
wordCountsCollected = (wordsRDD
    .< FILL IN >
    ...
    .< FILL IN >                       
    .collect())
```

<!-- grade_id: cell-efb14ccdee6629d8, points: 0, grade: false, solution: true -->
```python
wordCountsCollected = (wordsRDD
    .map(lambda word: (word, 1))  # Pair each word with 1
    .reduceByKey(lambda x, y: x + y)  # Count word occurrences
    .collect())  # Collect the results

# Print the results
for (word, count) in wordCountsCollected:
    print('{}: {}'.format(word, count))
```

Output:
```
lion: 1
tiger: 2
cat: 2
elephant: 1
rat: 2
```

<!-- grade_id: cell-856e3903dc0f06aa, points: 0.25, grade: true, solution: false -->
```python
# TEST exercise 2d
assert sorted(wordCountsCollected)==[('cat', 2), ('elephant', 1), ('lion', 1), ('rat', 2), ('tiger', 2)]
```

```python

```

## Part 3: Trobar les paraules individuals i la seva freqüència d'aparició mitjana

### (3a) Paraules úniques

Es demana obtenir el nombre de paraules úniques en 'wordsRDD'. Pots utilitzar altres RDDs que hagis creat en aquesta pràctica si et resulta més senzill.

**HINT (Copy & Paste and Complete):** 
```python
uniqueWords = < FILL IN >
```

<!-- grade_id: cell-16348ea08b49c4a8, points: 0, grade: false, solution: true -->
```python
# Get the number of unique words
uniqueWords = wordsRDD.distinct().count()

#Print the result
print(uniqueWords)
```

Output:
```
5
```

<!-- grade_id: cell-18114a75b3dc89f9, points: 0.4, grade: true, solution: false -->
```python
# TEST exercise 3a
assert uniqueWords== 5
```

### (3b) Calcular la freqüència mitjana d'aparició usant 'reduce()'

Troba la freqüència mitjana d'aparició de cada paraula en el RDD 'wordCounts'.

Es demana utilitzar l'acció 'reduce()' per a sumar els comptatges en 'wordCounts' i llavors dividir pel nombre de paraules úniques. Per a realitzar això primer aplica un 'map()' al pair RDD 'wordCounts', que està format per tuplas amb el format (key, value), per a convertir-ho en un RDD de valors.

**HINT (Copy & Paste and Complete):** 

```python
from operator import add
totalCount = (wordCounts
    .map(< FILL IN >)
    .reduce(< FILL IN >))
average = < FILL IN > / float(< FILL IN >)
```

NOTA: en cas de resoldre el problema seguint el procés anterior l'exercici es valorarà en 0.4 punts. Si es troba una solució que només faci servir una sola instrucció per a obtenir el resultat es valorarà en 0.6 punts. És a dir, sense haver d'aplicar la divisió en una segona instrucció.

<!-- grade_id: cell-9066ac11760c53f6, points: 0, grade: false, solution: true -->
```python
from operator import add

totalCount = (wordCounts
    .map(lambda pair: pair[1])
    .reduce(add))

average = totalCount / float(uniqueWords)

print(average)
```

Output:
```
1.6
```

<!-- grade_id: cell-074da01b339d1d9d, points: 0.4, grade: true, solution: false -->
```python
# TEST exercise 3b
assert round(average, 2)==1.6, 'incorrect value of average'
```

## Part 4: Aplicar les funcionalitats desenvolupades a un arxiu de text

L'objectiu que es proposa és construir una funció 'wordCount', capaç de treballar amb dades reals, que solen presenten problemes com l'ús de majúscules o minúscules, puntuació, accents, etc. Posteriorment, es demana carregar les dades de la nostra font de dades i finalment, calcular els contadors de paraules sobre les dades processades.

### (4a) Funció 'wordCount'

Es demana definir una funció per al comptatge d'aquestes paraules. S'aconsella reutilitzar les tècniques que has vist en els apartats anteriors d'aquesta pràctica. Aquesta funció, ha de prendre un RDD que contingui una llista de paraules, i retornar un pair RDD que contingui totes les paraules amb els seus corresponents comptatges.

<!-- grade_id: cell-90cb9e33943318d1, points: 0, grade: false, solution: true -->
```python
def wordCount(wordListRDD):
    #Convert all words to lowercase
    wordListRDD = wordListRDD.map(lambda word: word.lower())
    
    # Map words to (word, 1) pairs
    wordPairs = wordListRDD.map(lambda word: (word, 1))
    
    # Reduce by key to count word occurrences
    wordCounts = wordPairs.reduceByKey(lambda x, y: x + y)
    
    return wordCounts

#Apply function and show the result
countedWords = wordCount(wordsRDD).collect()
print(countedWords)
```

Output:
```
[('lion', 1), ('tiger', 2), ('cat', 2), ('elephant', 1), ('rat', 2)]
```

<!-- grade_id: cell-388aedb5b80c8820, points: 0.4, grade: true, solution: false -->
```python
# TEST exercise 4a
assert sorted(wordCount(wordsRDD).collect())==[('cat', 2), ('elephant', 1), ('lion', 1), ('rat', 2), ('tiger', 2)]
```

### (4b) Majúscules i puntuació

Els fitxers amb dades reals són molt més complexos que els que hem estat usant en aquesta activitat. Alguns dels problemes a tractar són:
+ Les paraules han de contar-se independentment de si contenen majúscules o minúscules (per exemple, Spark i spark haurien de contar com la mateixa paraula).
+ Tots els signes de puntuació han d'eliminar-se.
+ Qualsevol espai al principi o al final de la paraula ha d'eliminar-se.

Es demana definir la funció 'removePunctuation' que converteixi tot el text a minúscules, elimini els signes de puntuació i els espais al principi i fi de cada paraula. Fes servir el mòdul de Python [re](https://docs.python.org/2/library/re.html) per a eliminar qualsevol caràcter que no sigui una lletra, un número o un espai.

Si no estàs familiaritzat amb les expressions regulars s'aconsella revisar [aquest tutorial](https://developers.google.com/edu/python/regular-expressions). Alternativament, [aquesta web](https://regex101.com/#python) és de gran ajuda per a debugar les teves expressions regulars.

**Hints**

1. Usa la función [re.sub()](https://docs.python.org/2.7/library/re.html#re.sub).
2. Para nuestros propósitos, "puntuación" significa no alfabético, numérico, o espacio. La expresión regular que define estos carácteres es: `[^A-Za-z\s\d]`
3. No usar `\W`, ya que retendrá los guiones bajos.

<!-- grade_id: cell-2f6ff0c9ac2172e4, points: 0, grade: false, solution: true -->
```python
import re

def removePunctuation(text):
    # Regular expression to remove the non alphabetic characters and spaces
    cleanedText = re.sub(r'[^a-zA-Z0-9\s]', '', text)
    
    # Convert to lowercase and remove whitespaces
    cleanedText = cleanedText.lower().strip()
    
    return cleanedText

print(removePunctuation('Hi, you!'))
print(removePunctuation(' No under_score!'))
print(removePunctuation(' *      Remove punctuation then spaces  * '))
```

Output:
```
hi you
no underscore
remove punctuation then spaces
```

<!-- grade_id: cell-97a96198085b67f0, points: 0.4, grade: true, solution: false -->
```python
# TEST exercise 4b
assert removePunctuation(" The Elephant's 4 cats. ") == 'the elephants 4 cats'
```

### (4c) Carregar un fitxer de text

A partir d'aquest punt farem servir el text ja esmentat Lorem Ipsum generat per a la pràctica. Per a carregar un fitxer de text a un RDD, farem servir el mètode 'SparkContext.textFile()'. Una vegada carregat el fitxer es demana netejar les dades mitjançant 'removePunctuation()' dins d'una transformació 'map()'. Atès que el fitxer és bastant gran, farem servir 'take(15)', de manera que tan sols imprimirem per pantalla les 15 primeres línies.

Es demana revisar i executar la següent cel·la.

```python
# Just run the code
import os.path

fileName = os.path.join('/aula_M2.858/data/', 'LoremIpsum.txt')
loremRDD = sc.textFile(fileName, 8).map(removePunctuation).filter(lambda x: len(x)>0)
loremRDD.take(2)
```

Output:
```
['aut minima deleniti et autem minus illo esse dolores eligendi corrupti dolore minima nostrum eos nobis nam nihil aspernatur nam ut quae sint laborum ut dolores error possimus aperiam consequatur',
 'pariatur sed quo non itaque qui pariatur saepe ad quis consequatur nihil iste molestias et eos ut expedita vel reiciendis dolorem enim doloribus quam architecto aperiam']
```

```python
# Just run the code
import os.path

fileName = os.path.join('/aula_M2.858/data/', 'LoremIpsum.txt')
loremRDD = sc.textFile(fileName, 8).map(removePunctuation).filter(lambda x: len(x)>0)
loremRDD.take(15)
```

Output:
```
['aut minima deleniti et autem minus illo esse dolores eligendi corrupti dolore minima nostrum eos nobis nam nihil aspernatur nam ut quae sint laborum ut dolores error possimus aperiam consequatur',
 'pariatur sed quo non itaque qui pariatur saepe ad quis consequatur nihil iste molestias et eos ut expedita vel reiciendis dolorem enim doloribus quam architecto aperiam',
 'sed repudiandae pariatur similique est aut sequi animi in aperiam enim ipsa enim dolorem inventore aut quo odio in consequatur et',
 'aspernatur ad esse et aliquid itaque dolores rerum quia commodi explicabo non magnam nostrum consectetur non sint eum nulla et aut quis doloribus itaque nulla molestiae quis est est quo facilis incidunt a ipsa in itaque sed aut nobis facere dignissimos atque unde cum ea vero',
 'tenetur vel quod voluptatum laudantium dolores neque aut est modi qui aperiam itaque aperiam quae ratione doloremque aut delectus quas qui',
 'qui placeat vel ipsam praesentium sint recusandae dicta minus praesen
```

### (4d) Extreure les paraules de les línies

Abans de poder usar la funció 'wordcount()', hem de solucionar dos problemes amb el format del RDD:
+ El primer problema és que necessitem obtenir les paraules que conté cada línia. Això ho solucionarem en l'apartat (4d).
+ El segon problema és que necessitem filtrar les línies completament buides. Això ho solucionarem en l'apartat (4e).

Per a aplicar una transformació que divideixi cada element del RDD pels seus espais, es proposa aplicar la funció incorporada en els strings de Python [split()](https://docs.python.org/2/library/string.html#string.split). Cura que a primera vista pot semblar que la funció necessària és una transformacion 'map()', però si analitzes el resultat de la funció 'split()' t'adonaràs que aquesta no és l'opció correcta. Es demana a l'alumne revisar altres funcions afins a 'map()' en la documentació, que puguin ajustar-se millor a la tasca requerida.

> Nota:
> * No facis servir la implementació estàndard del 'split()', passa-li un valor de separador. Per exemple, per a dividir 'line' per comes, faries servir 'line.split(',')'.

**HINT (Copy & Paste and Complete):** 
```python
loremWordsRDD = loremRDD.< FILL IN >
```

<!-- grade_id: cell-62026cde37a4b3d5, points: 0, grade: false, solution: true -->
```python
loremWordsRDD = loremRDD.flatMap(lambda line: line.split())
loremWordsCount = loremWordsRDD.count()

print(loremWordsCount)
```

Output:
```
29249699
```

<!-- grade_id: cell-5eefb61d132b0437, points: 0.4, grade: true, solution: false -->
```python
# TEST Words from lines (4d)
# This test allows for leading spaces to be removed either before or after
# punctuation is removed.
assert loremWordsCount == 29249699
assert loremWordsRDD.top(5)==['voluptatum', 'voluptatum', 'voluptatum', 'voluptatum', 'voluptatum']
```

### (4e) Millorant l'eficiència mitjançant persistència

Una vegada que hem comptat el nombre de paraules amb el mètode 'count()', ens centrarem en fer el codi més eficient.

Per a reduir els temps d'execució, utilitzarem la persistència d'un RDD 'persist()', que principalment s'utilitza quan s'efectuen múltiples operacions repetides sobre un RDD. Quan es fa servir la persistència, cada node conserva la partició del RDD que opera en la memòria.

El bon ús de la persistència d'un RDD pot arribar a augmentar el rendiment de les aplicacions Spark en més de 10 vegades. Per a algorismes iteratius i aplicacions interactives ràpides, la persistència és de vital importància.

Existeixen diferents nivells de persistència, depenent de si s'emmagatzema el RDD en memòria o en disc. En aquest exemple farem servir el nivell de persistència per defecte MEMORY_ONLY.

Mitjançant les següents cel·les podrem veure els avantatges de la persistència. Per a poder quantificar la millora del rendiment ens ajudarem de la funció '%%time', que imprimeix el CPU.-time i el wall clock time d'execució per a la cel·la completa.

> Nota:
> * Wall clock time: temps total real que porta executar una tasca (CPU time, I/O time i temps de comunicació entre màquines si les dades estan repartides entre diverses)
>* CPU time: temps total que està la CPU executant una tasca

Seguidament, farem una sèrie d'execucions sobre un RDD utilitzant el mètode 'count()' encara que podríem fer servir diferents mètodes o combinacions de mètodes, però ho farem així per simplicitat.

1. Contem el nombre de paraules i veiem el temps invertit en l'execució

```python
%%time
loremWordsRDD.count()
```

Output:
```
CPU times: user 12 ms, sys: 8 ms, total: 20 ms
Wall time: 16.1 s
29249699
```

2. Es demada fer un `persist()` sobre el RDD `loremWordsRDD`
**HINT (Copy & Paste and Complete):** 
```python 
loremWordsRDD.< FILL IN >
```

<!-- grade_id: cell-842622900cd7989e, points: 0.1, grade: true, solution: true -->
```python
loremWordsRDD = loremWordsRDD.persist()
```

3. Tornem a comptar el nombre de paraules després de fer el `persist()` i observem que el temps d'execució és major que abans de fer el `persist()` ja que ha d'emmagatzemar els resultats en els nodes per al seu posterior ús.

```python
%%time
loremWordsRDD.count()
```

Output:
```
CPU times: user 0 ns, sys: 12 ms, total: 12 ms
Wall time: 22.9 s
29249699
```

4. En aquesta segona execució del `count()` i posteriors observarem que el temps d'execució és gairebé 10 vegades menor

```python
%%time
loremWordsRDD.count()
```

Output:
```
CPU times: user 12 ms, sys: 0 ns, total: 12 ms
Wall time: 4.63 s
29249699
```

5. Si volem des-persistir el RDD que hem persistit amb `persist()` farem servir `unpersist()`

**HINT (Copy & Paste and Complete):** 
```python
loremWordsRDD.< FILL IN >
```

<!-- grade_id: cell-04cf8ab9ef00be68, points: 0.1, grade: true, solution: true -->
```python
loremWordsRDD.unpersist()
```

Output:
```
PythonRDD[68] at RDD at PythonRDD.scala:53
```

6. Tornem a comptar el nombre de paraules amb el 'count()' i observem que tornem a tenir temps d'execució similars abans d'executar el 'persist()'

```python
%%time
loremWordsRDD.count()
```

Output:
```
CPU times: user 16 ms, sys: 4 ms, total: 20 ms
Wall time: 14.9 s
29249699
```

### (4f) Calcula paraules diferents

En aquest punt es demana comptar quantes paraules diferents conté el nostre text


**HINT (Copy & Paste and Complete):** 
```python
distintWordsMapRDD = loremWordsRDD.< FILL IN >
```

<!-- grade_id: cell-f9cb40f98d506e08, points: 0, grade: false, solution: true -->
```python
distinctWordsMapRDD = loremWordsRDD.map(lambda word: (word, 1))
distintWordsRDD = distinctWordsMapRDD.keys().distinct()
uniqueWordCount = distintWordsRDD.count()
print(uniqueWordCount)
```

Output:
```
182
```

<!-- grade_id: cell-a594923db15344f9, points: 0.3, grade: true, solution: false -->
```python
# TEST Remove empty elements (4f)
assert distintWordsRDD.count()== 182
```

### (4g) Compte les paraules

Es demana aplicar la funció `wordCount()` per a produir una llista amb els contadors de paraules al RDD 'loremWordsRDD' i mostrar les 15 paraules més comunes amb els seus comptatges. Podem fer servir l'acció 'takeOrdered()'; no obstant això, com els elements del RDD són tuplas, necessitem una funció especial que ordeni els parells de la forma correcta.

Fes servir les funcions `wordCount()` i `takeOrdered()` per a obtenir les 15 paraules més comunes juntament amb els seus comptatges.
**HINT (Copy & Paste and Complete):** 
```python
top15WordsAndCounts = wordCount(loremWordsRDD).< FILL IN >
```

<!-- grade_id: cell-d1578aa6bfcaee3f, points: 0, grade: false, solution: true -->
```python
# Get word counts using wordCount function
wordCounts = wordCount(loremWordsRDD)

# Get the 15 top words count
top15WordsAndCounts = wordCounts.takeOrdered(15, key=lambda x: -x[1])

# Display the result
for word, count in top15WordsAndCounts:
    print('{}: {}'.format(word, count))
```

Output:
```
et: 1291567
aut: 705322
ut: 704966
qui: 704703
est: 587782
voluptatem: 469634
quia: 469401
rerum: 353054
sed: 352873
omnis: 352653
consequatur: 352436
sit: 352396
non: 351959
voluptas: 351127
dolorem: 235822
```

<!-- grade_id: cell-55604f13d56cfa0a, points: 0.4, grade: true, solution: false -->
```python
# TEST Count the words (exercise 4g)
assert top15WordsAndCounts== [('et', 1291567), ('aut', 705322), 
                              ('ut', 704966), ('qui', 704703), 
                              ('est', 587782), ('voluptatem', 469634), 
                              ('quia', 469401), ('rerum', 353054), ('sed', 352873),
                              ('omnis', 352653), ('consequatur', 352436), 
                              ('sit', 352396), ('non', 351959), ('voluptas', 351127), 
                              ('dolorem', 235822)]
```

## Part 5: Calcular algunes estadístiques

Fent servir les mateixes tècniques que has aplicat en els exercicis anteriors respon a les següents preguntes:

- Quantes paraules diferents tenen exactament dos 's'?
- Com és la paraula de set lletres que més es repeteix? Quantes vegades apareix?
- Quantes paraules diferents tenen més consonants que vocals?

<!-- grade_id: cell-7c371ee95bd57d18, points: 0, grade: false, solution: true -->
```python
#¿Cuantas palabras distintas tienen exactamente dos 's'?
# TODO: Complete the cell with the appropriate code
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-072520fa34734220, points: 0.6, grade: true, solution: false -->
```python
# TEST Different words with exactly two 's' (5a)
assert loremWordsRDD_2s.count()== 7
```

<!-- grade_id: cell-079fd33957109dca, points: 0, grade: false, solution: true -->
```python
#¿Cual es la palabra de siete letas que más se repite?¿Cuantas veces aparece?
# TODO: Complete the cell with the appropriate code
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-b7efa700f3bef7bb, points: 0.6, grade: true, solution: false -->
```python
# TEST 7 letters most repeated word, how many times is repeated (5b)
assert distintWordsRDD_7_Number_Repetitions.count()== 23
```

<!-- grade_id: cell-a5272a8fe999e503, points: 0, grade: false, solution: true -->
```python
#¿Cuantas palabras distintas tienen más consonantes que vocales?
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ef4bc4b82edc1fe2, points: 0.8, grade: true, solution: false -->
```python
# TEST More consonants than vowels (5c)
assert moreConstWordsRDD.count()== 89
```

<h2 style='margin-left:18.0pt;text-indent:-18.0pt'>6<span style='font:7.0pt "Times New Roman"'>&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;
</span>Video resposta - RDDs</h2>

En aquest apartat s'haurà de respondre a les qüestions que es mostren a continuació mitjançant un vídeo resposta d'entre 5 i 10 minuts. Heu de realitzar una gravació en la qual se us vegi la cara responent amb les vostres pròpies paraules.

Heu de seguir els següents punts per a realitzar el vídeo:

1. Al principi del vídeo heu de dir el vostre nom i DNI mentre mostreu el DNI (que ha de veure's amb claredat i ser llegible).

2. Heu de respondre a les preguntes en l'ordre en el qual estan, comentant sempre al principi de cada resposta a què pregunta us esteu referint.

3. En el cas de no respondre a alguna pregunta, heu d'indicar el número de la pregunta que no responeu i el perquè.


    
<h3 style='margin-left:18.0pt;text-indent:-18.0pt'> Preguntes </h3>


Contesta a les següents preguntes en relació amb la utilització de Spark amb RDDs:

1 - Quina és la diferència entre accions i transformacions? (0.5 punts)

2 - Quins nivells de persistència es poden utilitzar? (1 punt)

3 - Què és el "Spark Driver" ? (0.5 punts)

4 - Quins sistemes de fitxers suporta Spark ? (0.5 punts)

5 - Quin "Clúster Managers" es poden utilitzar amb Spark. ? (0.5 punts)

<!-- grade_id: cell-174494fdb2664d21, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```